Analysis of NIH K99/R00 Awards and the Career Progression of Awardees

Author Response

The following is the authors’ response to the previous reviews.

Recommendations for the authors:

Reviewer #1 (Recommendations for the Authors):

The authors have addressed my recommendations in the previous review round in a satisfactory way. I only have one additional comment to the authors:

In the manuscript abstract lines 31-32, the author state that: "Using NIH data for the period 2006-2022, we report that ~230 K99 awards were made every year, representing ~$25 million annually."-- The "~$25 million" is under-stating the actual funds spent because this sum is just money spent on the first year of some k99s while the NIH is paying years 2,3,4 etc for others for k99 awards (~90% conversion rate to R00) awarded in previous years for a given year. The NIH is actually spending ~$230-$250 million a year on the k99 award mechanism in a given year. so the authors need to amend the stated amount in the manuscript.

Thank you for pointing this out. The reviewer is correct, that we had incorrectly only calculated the investment $ in new K99 awards made. We have corrected this in the revised manuscript. We appreciate your careful reading of our manuscript and the edits made based on your comments have improved the final version.

Reviewer #2 (Recommendations for the Authors):

Thank you for taking the time to revise this important work. I learned a lot reading this paper a second time, and appreciate the improvements you have made.

My only major thought while re-reading this is that I wish you all had written two papers! I see two themes in this work: one looking at faculty hiring networks from the Wapman et al. dataset, and another at K99/R00 conversions by institution, gender, and researcher mobility and its impact on subsequent funding success. After reading, I felt like I had many follow-up questions about both analyses, but it would be impractical for me to suggest all these follow-up analyses without making your paper unreasonably long.

Thank you for these comments. We agree that there are 2 general themes in this paper. While we feel that significantly expanding on both themes will be important in future research. Our hope is that this work continues to inspire others to critically examine funding practices and inequity in the same way that the work of Wapman, Pickett, etc. inspired the present work.

For example, regarding the results that more R00 are activated at different institutions, and that moving institutions improves subsequent funding success, I wonder: Do proportionally more women or men move institutions? Do proportionally more K99 awardees at less-funded places move for their R00, or less? The Cox proportional hazard models illustrate the impact of various characteristics on subsequent funding success, but they do not illustrate disparate impacts of mobility on different groups (if I am understanding them correctly). (You sort of dive into these questions in the very interesting subsection, "K99/R00 awardee self-hires are more common at institutions with top NIH funding." I wanted to read more!)

Thank you for these kind comments. These are fantastic follow-up questions. We do not feel that we can adequately address them within the present manuscript without potentially splitting it into 2 separate manuscripts. However, we may examine these in future analyses. We are particularly interested in examining additional aspects such as how the K99 MOSAIC funding mechanism may differ from the traditional K99 mechanism. Since the K99 MOSAIC mechanism is newer, there may not be enough K99 MOSAIC awards made for a thorough exploration.

As another example, for your analysis on faculty hiring networks, the prevalence of self-hiring amongst institutions and regions was one finding. However, this finding seems somewhat at odds with the previous takeaway about how researcher mobility improves subsequent funding success. Are institutions doing themselves a disfavor by hiring their own, then? I suspect there is more to say here about this pattern... maybe there are important differences between PhD institution and postdoc institution and its impact on hiring/subsequent funding success? Or is this a story about upward mobility into the top 25 well-funded NIH institutions?

Again, these are very insightful comments and follow-up questions. We hope to address these in potential future manuscripts. We also hope that others may become interested in finding answers to these questions by exploring our dataset as well as other publicly available datasets such as the Wapman et al. dataset.

I can completely understand how combining the faculty hiring network analysis with the K99/R00 conversions would seem like a natural fit, but I personally feel - emphasis on this being a personal opinion - that there would have been benefits to giving more space to the details of both analyses separately. Perhaps this is a "hindsight is 20/20" issue. Or an issue with the current times in which ones' brain can only hold so many main takeaways from a single body of work. (For example, I struggled to summarize your paper in my public review because I find so many takeaways important.)

I suppose this is all to say that I find your work important enough to warrant additional follow-up work! :)

Thank you for these very kind remarks. This work evolved over 8-10 months as evidenced by the updates to the biorXiv preprint. With unlimited time and foresight, it would probably be best to have separated the 2 themes into separate manuscripts and expanded both. Given current constraints, we plan to make some changes/updates to the present manuscript and hopefully include more in-depth analyses on each theme in future works. Thank you again for the thoughtful reading and critique of both our original manuscript and the revised version.

Minor comments/questions:

"K99 to R00 conversions are increasing in time"

Assuming I am interpreting the figures correctly, in my opinion, the most important takeaway is that the number of R00 awards have increased, but only for awardees moving to another institution. This key result, best illustrated by panels A and C of Figure 1, is buried in the long paragraph in this section. The organization of content in this section could be improved and more focused. Consider renaming this subsection to be more declarative: "K99 tR00 conversions have increased, but only for awardees moving to another institution."

This is a very concise interpretation of this data. We have edited the paragraph referenced by the reviewer, split it into 2 paragraphs, and changed the title to “K99 awardees increasingly move to other institutions for R00 awards from 2008 to 2022” and the final sentence to “Thus, the number of K99 to R00 conversions is consistent over time, but increasingly more R00 awardees have moved to other institutions since 2013”

Similarly, I personally found the current title of the subsection, "K99 to R00 conversions are increasing with time" is mildly confusing. An R00 award indicates a successful conversion, so why not simply call this an R00 award instead of saying K99-to-R00 conversion? Also, when I look at Figure 1B and exclude the conversion rates for 2007 and 2008 (because this is a 3 year rolling average), I see that conversion rates (or R00 awards) have remained stagnant. This comment is very much in-the-weeds and is mainly to do with clarity of language.

Thank you for these comments. We had “K99 to R00 conversion” to highlight the unique nature of this award mechanism that a person can only receive an R00 if they previously had a K99 award. Nevertheless, we have edited the text to “R00 awards” and “R00 awardees” to simplify things. We also want to note that we did not compute a 3-year rolling average. The function we used was: (X/(Y -1))x100 where X is the number of R00 awards made in a year and Y is the number of K99 awards made in a year. We did note an error in our calculation in the previous version of the manuscript. Previously, we included all R00 awards and K99 awards for each year from the NIH Reporter dataset; however, this is a flawed methodology. NIH reporter includes only extramural K99 award data and extramural R00 awards, but intramural K99 awardees can receive extramural R00 awards and thus are only included in the R00 dataset. There were 141 R00 awardees in our dataset from NIH Reporter that did not have K99 data, so we assume these are intramural K99 awards since it is required to have a K99 to be eligible for the R00 award. Since we do not know the awarding year for intramural K99 awardees or have data on intramural K99 awardees that fail to activate the R00 award (or stay internal at NIH), we have excluded these 141 R00 awardees. In the previous version, this mis-calculation exaggerated rolling conversion rate (we had correctly calculated the 78% total conversion rate). We re-analyzed our rolling conversion rate and found the average is 81.8% (excluding the first 2 years of the K99 program and the last 2 years).

This is a long explanation, but essentially, we overestimated the number of R00 awards which inadvertently increased the rolling conversion rate. We have corrected this and simplified the first 2 paragraphs of the Results section.

I was also mildly confused looking at Figure 1c. The caption says that the percentages represent the K99 awardees that stayed at the same institution for the R00 activation, but the percentages are next to the solid circles which the legend labels as "different institution." Perhaps another or different way to show this is a stacked bar chart, where one bar represents the percentage of R00 awards activated at the same institution and another bar represents the percentage of R00 awards activated at a different institution. The bars always add to 100% but the change in proportions illustrates that proportionally fewer awards are being made to those remaining at the same institution.

Great idea. We have included a stacked bar chart here. Since the stacked bar chart is percentages, we felt it was important to also show the total numbers so we still included the previous chart also but removed the percentage numbers from it. We also changed the departmental analysis to stacked bar charts. This shows the stark difference between 2008-2012 and 2013 onward. These changes were made in the revised Fig. 1.

Minor question: I would love to see Table 3 and Table 4 as a time-series. Has the proportion of recipients at various institution types changed with time?

This is a great suggestion and we felt it fit best in Figure 5, so we’ve added it there.

Table 3 is useful but only indirectly addresses my first "Recommendation to the Authors" from my previous review. I did some number crunching myself from the data provided. Assuming I did this correctly: If you're a K99 awardee at a private institute, you had a 76.3% change of getting an R00 compared to 80.4% for a K99 awardee at a public institution. If you're a K99 awardee at a top-funded institution, you had a 76.8% chance of R00 compared to 78.6% for a lower-funded institution. I would have liked to see more figures and tables to illustrate conversion rates by institution type in this way. Interestingly, to me, these data suggest that there are not enormous conversion rate differences by institution type (though looking at these now, I am confused at the 89% statistic in line 174 and where that comes form, since it is much higher than what I've calculated).

Thank you for this suggestion and these comments. Please see above where we describe how we incorrectly overestimated the 89% statistic. This has been corrected. As the reviewer suggested, we now show yearly percent of grants to specific institution types in the revised Figure 5. We agree with the reviewer that showing the conversion rate by institution type is interesting; however, it is fairly obvious from the new panels in Figure 5 that there is not much difference in conversion rate. Thus, to avoid crowding too many panels into the figure, we opted to keep the stacked bar plot.

Reviewer #3 (Recommendations for the Authors):

-One minor change to Figure 1C would be to switch the color coding for the lines so that they match with 1D whereby "same institution" would be white circles, or whatever the authors decide would be best for consistency since they are similar comparisons.

Thank you for this suggestion. We have corrected this to be consistent.

-Minor note for lines 459-461: I would suggest changing the wording to "intersectional inequalities" as it is not that a scientist's identities impact their careers as much as how those identities are positioned within an unequal opportunity structure and differentially treated that produce varying career trajectories and experiences of marginalization and cumulative (dis)advantages.

Thank you and we agree with you. We have made this correction.

-To carry forward a suggestion for the authors in my previous review, future research that more fully explores the research infrastructure of institutions for how top NIH funded institutions continue to be top funded institutions year after year could help clarify some of the career mobility and same/similar institution hiring found in the data. Rather than hand coding institutions for some of the infrastructure, the National Center for Education Statistics' Integrated Postsecondary Education Data System (IPEDS) has data on colleges and universities including whether they operate a hospital, have a medical degree, and many other interesting data about student and faculty demographics, institutional expenditures (including research budgets), and degrees awarded in different fields of study (undergrad and grad) that may be helpful to the authors as they continue their research stream in this area.

Thank you very much. We will look into this data set as we continue our investigations in this area.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Many postdoctoral fellows and scholars who hope to secure tenure-track faculty positions in the United States apply to the National Institutes of Health (NIH) for a Pathway to Independence Award. This award has two phases (K99 and R00) and provides funding for up to five years. Using NIH data for the period 2006-2022, we report that ∼230 K99 awards were made every year, representing ∼$25 million annually. About 40% of K99 awardees were women and ∼89% of K99 awardees went on to receive an R00 award annually. Institutions with the most NIH funding produced the most recipients of K99 awards and recruited the most recipients of R00 awards. The time between a researcher starting an R00 award and receiving a major NIH award (such as an R01) ranged between 4.6 and 7.4 years, and was significantly longer for women, for those who remained at their home institution, and for those hired by an institution that was not one of the 25 institutions with the most NIH funding. Shockingly, there has yet to be a K99 awardee at a historically Black college or university. We go on to show how K99 awardees flow to faculty positions, and to identify various factors that influence the future success of individual researchers and, therefore, also influence the composition of biomedical faculty at universities in the US.

Revised version: Many postdoctoral fellows and scholars who hope to secure tenure-track faculty positions in the United States apply to the National Institutes of Health (NIH) for a Pathway to Independence Award. This award has two phases (K99 and R00) and provides funding for up to five years. Using NIH data for the period 2006-2022, we report that ∼230 K99 awards were made every year, representing up to ∼$250 million annual investment. About 40% of K99 awardees were women and ∼89% of K99 awardees went on to receive an R00 award annually. Institutions with the most NIH funding produced the most recipients of K99 awards and recruited the most recipients of R00 awards. The time between a researcher starting an R00 award and receiving a major NIH award (such as an R01) ranged between 4.6 and 7.4 years, and was significantly longer for women, for those who remained at their home institution, and for those hired by an institution that was not one of the 25 institutions with the most NIH funding. Shockingly, there has yet to be a K99 awardee at a historically Black college or university. We go on to show how K99 awardees flow to faculty positions, and to identify various factors that influence the future success of individual researchers and, therefore, also influence the composition of biomedical faculty at universities in the US.

Introduction

Postdoctoral research fellows are a driving force of the academic biomedical research enterprise. They typically have doctorate (i.e., PhD) or medical degrees (i.e., MD or DVM) which means they have completed a significant amount of education, training, and publishing.1–4 Postdoctoral fellows aspiring to continue in academic research typically seek tenure-track faculty appointments.4 Despite recent trends of new doctoral degree holders seeking private sector positions5, tenure-track academic faculty positions are highly coveted and competitive due to the increasing number of PhD graduates but stagnant number of tenure-track faculty positions.3,6 Obtaining such positions requires consistent access to opportunities for success such as attending doctoral programs at perceived prestigious institutions, publishing manuscripts at the undergraduate, graduate, and postdoctoral levels, and receiving extramural grant funding throughout one’s career.7–10 University prestige has significant impacts on resources11, paper acceptance rates12, citations13, and awards.14 There is also significant bias against women, LGBTQIA+, and other systematically marginalized groups at all levels, which results in underrepresentation in STEM fields and fewer R01 awards compared to white/men counterparts.15–18 To combat these biases, the NIH developed funding mechanisms, such as the K99 MOSAIC and K12 IRACDA awards, to promote scientists from diverse backgrounds into tenure-track faculty positions. In addition, other non-federal grant mechanisms exist, such as the Gilliam Fellowships for Advanced Study from the Howard Hughes Medical Institute which began in 2004. The success of these programs at addressing bias or improving equity for systematically marginalized groups in the tenure-track ranks is currently unclear because the programs are either too new (MOSAIC) or data are not publicly available (HHMI Gilliam Fellowships).

Within the United States, one of the most lucrative award mechanisms to facilitate the transition from postdoctoral fellow to tenure-track faculty member in the biomedical sciences is the NIH Pathway to Independence Award (K99/R00).19 Launching in 2006, this award provides salary support for postdoctoral fellows for two years during the K99 phase and $250,000/yr for 3 years during the R00 phase.20,21 The candidate benefits directly through funding, visibility, and protected time.1,21 Having a K99/R00 award also demonstrates to hiring institutions that the candidate has a track record of receiving extramural funding which is a requirement for future success in academic research careers.22 Roughly 89% of K99 awardees received R00 funding on a rolling yearly basis, which indicates having obtained a faculty position although factors such as publications, presentations, interpersonal skills, choice of advisor, and others also contribute to securing faculty positions. Of those that secure faculty positions, many factors influence postdoctoral researchers in their decision of which position to choose including career-oriented factors like department, support, colleagues, and personal factors such as family and location. About 50% of R00 awardees subsequently attain R01 funding23 which is a higher success rate than other NIH career development grants such as K01, K08, and K2319,24, although having any K-award tends to increase future funding success.25 The specific and cumulative factors that influence future NIH funding success for R00 awardees has not been examined.

Individuals from marginalized communities are typically at a disadvantage when applying for research grants such as R01s. For example, men typically experience 2-3% greater funding success rates; however, in 2019 women had an advantage in funding success rates for NIH research project grants (21% versus 20%).26 Men also receive an average of $35,000-$45,000 more NIH funding per grant than women or other gender identities,17,26,27 however, some report no difference.28 Women have been reported to have a disadvantage in transitioning from K-awards to R-awards from the NIH.29 Racial and ethnic minorities are also at a disadvantage in receiving R01 funding compared to white counterparts.30,31 The institution itself also influences R01 funding success. Here, we explore the conversion of K99 to R00 awards by year, gender identity, and institutions (other demographic data are not publicly available) to demonstrate the flow of K99/R00 awardees from postdoc to faculty positions. We also determine whether these influence future NIH funding success for K99/R00 awardees. These findings have the potential to influence how and where career development awards are made, how potential candidates are supported, and how grant reviewing practices may be changed to be more equitable and inclusive.

Methods

Data acquisition and analysis

Raw K99/R00 data examined in this manuscript is publicly available from the NIH Reporter (https://reporter.nih.gov/). Filtered data and annotated data used to generate Figures 1-5 and Tables 1 & 2 are available in the Supplemental Tables. Names have been removed; however, the corresponding author will provide the identified data upon reasonable request. All data for K99 awardees, R00 awardees, and matched K99/R00 awardees is included in Supplemental Tables 1-3. The data used to generate Figure 6 and Supplemental Figures S1 and S2 were generated by Wapman et al. and are available on Zenodo at https://doi.org/10.5281/zenodo.6941651.15 K99 and R00 data were downloaded from NIH Reporter using the “Advanced Search” function to include only the year range (2006-2022) and only the grant of interest (i.e., K99 or R00). Duplicate values containing additional awards for the same contact PI (i.e., supplemental awards or multiple years of the same award) were removed. The latest available fiscal year for the K99 award was used and the earliest available fiscal year for the R00 was used. The K99 and R00 institution were aligned and confirmed that the contact PI name was identical before removing the contact PI name from the dataset. The institutions were analyzed using Microsoft Excel and Graphpad Prism (version 9.4.1). Data on doctoral degree university and subsequent faculty position are from Wapman et. al, 2022 and are available at https://larremorelab.github.io/us-faculty/.15 We included the following fields from this dataset: Anatomy, Animal Sciences, Biochemistry, Biological Sciences, Biomedical Engineering, Biophysics, Biostatistics, Environmental Health Sciences, Exercise Science, Kinesiology, Rehab, Health, Health, Physical Education, Recreation, Human Development and Family Sciences, Immunology, Microbiology, Molecular Biology, Neuroscience, Nursing, Nutrition Sciences, Pathology, Pharmaceutical Sciences, Pharmacology, Pharmacy, Physiology, Psychology, Public Health, and Veterinary Medical Sciences. We examined the Historically Black Colleges and Universities from this dataset. The list of K99 institutions and R00 institutions and HBCU institutions were plotted as Sankey diagrams using the “gvisSankey()” function in the “googleVis” package in R (version 2022.07.1, Build 554). Interactive versions of the K99 to R00 Sankey plot is available at: https://k99tor00.shinyapps.io/K99-R00_Sankey/ with data available here: https://zenodo.org/records/10005359 and code available here: https://github.com/chsolis/K99toR00SankeyNetwork2007-2022. Interactive versions of the HBCU Sankey plots (Fig. 7) are available at: https://dantyrr.github.io/html/. The K99 and R00 location maps were generated on ObservableHQ using a map layer and overlaying heatmap dot plot by latitude and longitude coordinates from the NIH Reporter dataset. Data on R01, DP2, and R35 awards were exported from NIH Reporter using the list of K99 to R00 transfers for each year from 2007 to 2023 and the time from each candidate’s first R00 budget start date to R01, DP2, and R35 budget start date was calculated.

Earlier version: Rate of K99 and R00 awards and conversion by year. A) Total extramural K99 awards and R00 awards made by year. B) Annual rolling conversion rate of R00 award activation compared to the number of K99 awards activated in the prior year. Average rolling retention rate (89%) from 2008-2020 shown in dashed line. C) Total number of extramural K99 to R00 conversions per year at the same institution or different institution. For each year, numbers represent the percentage of K99 awardees that stayed at the same institution for the R00 activation. D) Within the K99 to R00 conversions at the same institution, the number of K99 to R00 conversions either in the same department or different department, by year. For each year, numbers represent the percentage of K99 awardees that stayed in the same department.

Revised version: Rate of K99 and R00 awards and conversion by year. A) Total extramural K99 awards and R00 awards made by year. B) Annual rolling conversion rate of R00 award activation compared to the number of K99 awards activated in the prior year The formula is: (X/(Y -1))x100 where X is the number of R00 awards made in a year and Y is the number of K99 awards made in a year. Average rolling retention rate (81.8%) from 2008-2022 shown in dashed line. C) Annual rolling conversion rate of R00 award activation compared to the number of K99 awards activated in the prior year stratified by whether the R00 was activated at the same or different institution as the K99 award. Average retention rates for R00 at different institution (93.9%) and the same institution (69.4%) are shown in dashed lines. D) Total number of extramural R00 awards per year at the same institution or different institution. E) Annual stacked bar plots showing the percentage of R00 awards and whether they were at the same institution or different institution. F) Within the K99 to R00 conversions at the same institution, annual stacked bar plots showing the percentage of R00 awards and whether they were at the same department or different department.

Institutional classification

Institutions were classified as either “highest NIH funding” or “lower NIH funding” based on 2022 NIH funding levels and included the 25 institutions with the most NIH funding. The institutions included in this list in order from most funding to least are: Johns Hopkins University, University of California, San Francisco, University of Pittsburgh at Pittsburgh, Duke University, University of Pennsylvania, Stanford University, University of Michigan at Ann Arbor, Washington University, Columbia University Health Sciences, University of California, San Diego, University of California Los Angeles, Yale University, University of Washington, Univ of North Carolina Chapel Hill, Massachusetts General Hospital, Emory University, Icahn School of Medicine at Mount Sinai, University of Minnesota, Fred Hutchinson Cancer Center, Northwestern University at Chicago, Brigham and Women’s Hospital, New York University School of Medicine, Vanderbilt University Medical Center, and University of Wisconsin-Madison. Leidos Biomedical Research, Inc. is also included in this list but does not have any K99 or R00 awardees.

Binary name-to-gender classification

Name-based ascription tools for gender and for race/ethnicity are inherently biased and these biases occur unevenly across groups.32 A name-to-gender assignment tool (GenderAPI, https://gender-api.com/) was used to assign a binary gender (man/woman) to K99/R00 awardees. GenderAPI was chosen based on its low rate of mis-classifications and non-classifications33. It assigns gender identity as either men, women, or unknown. Because the terms male and female typically refer to biological sex, the terms man/men or woman/women are utilized here. However, it is important to note that gender identity is a social construct that is non-binary and non-static. Individuals identifying as transgender, non-binary, or gender non-conforming, or whose gender identity has changed since the time of receiving the K99 or R00 award are at risk of being misgendered. We acknowledge this remains a major limitation of this type of name-to-gender methodology and we hope it changes in the future to be more inclusive.

Results

Added in the revised version: K99 awardees increasingly move to other institutions for R00 awards from 2008 to 2022

Removed in the revised version: K99 to R00 conversions are increasing in time

Earlier version: K99 awards provide funding for 2 years and awardees generally can accept faculty positions within the second year. From the total of 3475 extramural K99s awarded from 2006 through 2021 (Fig. 1A), 2703 extramural K99 awardees that converted to R00 awards by the end of 2022 (Fig. 1B). Thus, a total of 78% of K99 awardees converted to the R00 phase but on a rolling basis excluding the first 2 years and last 2 years of the program, ∼89% of K99 awardees go on to receive the R00 award. From 2007-2022, there were a total 2847 R00 awards (Fig. 1A), and thus we conclude that the remaining 144 R00 awardees with no matching K99 data are likely intramural NIH K99 awardees. There is no public data on intramural K99 awards in NIH Reporter. This means that there are 772 scientists with K99 awards did not get R00s. 468 of these K99 awards were from 2020 or later and may still convert to the R00 phase; however, 304 were from 2019 or earlier, which suggests they have not and will not convert to the R00 phase. Where those scientists went and why they did not convert the K99 grant into an R00 is unknown. From 2008 through 2021, there were an average of 234 K99 awards/year, and from 2009 to 2022, an average of 198 K99 awardees converted to R00 awards/year (Fig. 1B). In the early years of the K99/R00 program, around 30% of R00 awardees had been hired at the same institution. From 2012 to 2016, the proportion of awardees that stayed at the same institution declined to about 20% where it has remained consistent (Fig. 1B). In the beginning of the K99/R00 program, nearly all the awardees that stayed at the same institution also stayed in the same department; however, from 2013 onward a greater number moved to new departments (Fig. 1C). This is in line with what Pickett reports that ∼20% of K99 recipients received their first R01 award at the same institution where they had a K99 award.23 In general, the majority of K99 awardees in recent years move to new institutions to activate the R00 phase and begin faculty positions.

Revised version: K99 awards provide funding for 2 years and awardees generally can accept faculty positions within the second year and convert the K99 award to an R00 award. From the total of 3456 extramural K99s awarded from 2006 through 2021 (Fig. 1A), 2842 extramural R00 awards have been made by the end of 2022 (Fig. 1A). 141 of the R00 awardees do not have K99 data in the NIH Reporter database and thus must be intramural K99 awardees (i.e., awardees at NIH research laboratories). Since we do not have data on intramural K99 awardees who do not convert to the R00 award, we will exclude these 141 R00 awardees. There are 755 extramural K99 awardees that did not receive R00 awards. 426 of these received K99 awards in 2020 or after and may still be in the K99 phase; however, 329 were from 2019 or earlier, which suggests they have not and are no longer eligible for R00 awards. Where those awardees went and why they did not receive an R00 is unknown.

Added in the revised version: From 2008 through 2021, there were an average of 234 extramural K99 awards made per year, and from 2009 to 2022, an average of 198 R00 awards per year. 2701 extramural K99 awardees activated R00 awards, for a total of 78% conversion. On a rolling basis, excluding 2006-2007 (first 2 years), 81.8% of K99 awardees go on to receive an R00 award (Fig. 1B). Until 2015, the rolling conversion rate was similar among R00 awardees who were at different institutions than the K99 award and R00 awardees at the same institution (Fig. 1C). Starting in 2015, the activation rate of R00 awardees at the same institution decreased and has been lower than R00 awardees at different institutions since then (Fig. 1C). The average rolling conversion rate for R00 awardees at the same institution was 69.4% whereas it was 93.9% for R00 awardees at different institutions. In the first 4 years of the K99/R00 program, around 30% of R00 awardees had been hired at the same institution; however, from 2013 onward, the proportion of awardees that stayed at the same institution declined to ∼20% and has remained consistent (Fig. 1D-E). The number of R00 awardees at the same institution was consistent from 2010-2022 but the number of R00 awardees at different institutions has increased over the same period (Fig. 1D-E). In the beginning of the K99/R00 program, nearly all the awardees that stayed at the same institution stayed in the same department; however, from 2013 onward a greater number moved to new departments (Fig. 1F). This is in line with what Pickett reports that ∼20% of K99 recipients received their first R01 award at the same institution where they had a K99 award.23 Thus, the number of K99 to R00 conversions is consistent over time, but increasingly more R00 awardees have moved to other institutions since 2013.

Geographic localization of K99 awards are at a more concentrated set of institutions whereas R00 awards are more dispersed

K99 awards have been granted to individuals at 256 different institutions and R00 awardees represented 357 unique institutions. Strikingly, 54% of all K99 awardees were at 10% of all institutions, while 46% of all R00 awardees were at 10% of institutions. K99 and R00 awards were both most frequent in parts of the country containing a high concentration of research universities (e.g. the Bay Area in California and Boston Area in Massachusetts) (Fig. 2). Three institutions had more than 75 K99 awardees through the end of 2022 (Stanford [237], University of California San Francisco [116], and Massachusetts General Hospital [93]), and 73 institutions had just one, with the median being four awards per institution. Remarkably, Stanford had more awards than the 134 institutions with the fewest number of awards (Figure 2). R00 awards were dispersed across more institutions than K99 awards, with the maximum number being 60: three institutions had more than 50 (Massachusetts General Hospital [60], University of Pittsburgh [58], and University of Michigan [54]), and 116 institutions had just one (Figure 2). K99 awards are more concentrated in East and West coast states than R00 awards (73% vs 55%). California (n=491), Massachusetts (n=348), and New York (n=121) had the largest proportion of K99 relative to R00 awards, and Texas (n=48), Ohio (n=46), Michigan (n=33), and Arizona (n=31) received the greatest number of R00 awardees relative to K99 awards. These data indicate that while more K99 awards are physically located on the East and West coast universities, the R00 awardees are distributed more evenly throughout the geographic location of the United States.

Cartographic representation of the number of K99 and R00 awards by institutional location and state from 2007-2008 through 2021-2022. Note no individual award or grant is counted more than 1 time.

Institutions with the most NIH funding tend to hire K99/R00 awardees from other institutions with the most funding

We next analyzed the flow of awardees from K99 institution to R00 institution (n=2703 individuals) (Fig. 3). We coded the 25 institutions that received the most NIH funding in the year 2022 as “1” and the rest as “2”; however, Leidos Biomedical Research, Inc. is included in the highest funded institutions but has no K99 or R00 awardees. The institutions that produce the highest number of K99 awardees also tend to hire the most R00-awardees and also tend to receive the most NIH funding (Fig. 3). Most K99 awardees attain faculty positions at private institutions (62%) compared to public institutions (38%). Private institutions with the highest NIH funding generally hire R00 faculty from other private institutions whether they are in the top NIH funding category (31.2%) or not (46.9%). The public institutions with the highest NIH funding and private institutions that do not have the highest NIH funding also tend to hire R00 faculty more from private institutions categorized as both top NIH funding (20.3% and 22.5%, respectively) and lower NIH funding (27.2% and 41.3% respectively). Notably, the public and private institutions with the largest NIH funding portfolios and other private institutions have hired only 26% of K99 awardees from public institutions with smaller NIH portfolios. In contrast, these public institutions with smaller NIH funding portfolios have hired 74% of all K99 awardees from similar institutions (Table 1).

Sankey diagram of all successful K99 to R00 award transfers from 2007-2022 with the K99 institution on the left and R00 institution on the right. An interactive version of this figure can be found at: Interactive Sankey diagram can be found at: https://k99tor00.shinyapps.io/K99-R00_Sankey/

K99/R00 awardee self-hires are more common at institutions with the top NIH funding

We also demonstrate the number of self-hires of K99/R00 recipients from within institutions. A mean of 27% of K99 awardees remain at their home institution during the R00 phase; however, this differs based on institutional type and NIH funding portfolio size. At private institutions with large NIH funding portfolios, 36% of those stay at the same institution for the R00 phase. Fewer individuals stay at the same institution from public institutions with large NIH funding portfolios (30.9%), private smaller NIH funding portfolios (23.3%), and public smaller NIH funding portfolios (18.5%). 70% of the R00 awardees recruited by Massachusetts General Hospital were self-hires. Of the remaining 17 individuals, 8 were from other Harvard-associated institutions (Brigham and Women’s Hospital, Dana Farber Cancer Institute, Broad Institute). There are other institutions that have higher relative proportions of K99 awardees remain for the R00 phase (i.e., University of Pittsburgh, Johns Hopkins University, and Brigham and Women’s Hospital). In contrast, some institutions recruit more of their R00 awardees from external institutions (i.e., Northwestern University, University of Texas Southwestern, University of Utah, and Yale University).

More K99 awardees are men than women

We performed a name-to-gender classification using binary terms of man and women. This type of classification is inherently biased unevenly across demographics and groups (please see Methods).32 Using this methodology, we found that men make up the majority of K99 (n=2028, 58%) and R00 (n=1655, 58%) awardees. The same percentages of men and women K99 awardees convert their K99 awards to R00 awards (Table 2). Since the beginning of the K99 award mechanism in 2006, men have received more K99 awards than women each year; however, the percentage of women receiving K99 awards has risen slightly in the most recent decade (Fig. 4).

Number of K99 awards per year to either men or women (left Y axis) and the percentage of K99 awardees that are women (right Y axis).

K99 and R00 awards are concentrated within the highest funded institutions

Earlier version: The majority of K99 awardees were at private institutions with 28% from the top 25 highest funded institutions and 35% from other private institutions. The remaining 37% of K99 awardees were from public institutions (Fig. 5 and Table 4). In contrast, the majority of R00 awardees were at public institutions outside of the top 25 highest funded (38%). The second largest group of R00 awardees were at private institutions outside of the top 25 highest funded (29%) (Fig. 5 and Table 4). To further put this in context, we examined the composition of institutions represented by both K99 awards and R00 awards. We found that 45% of K99 awards were made to the top 25 highest funded institutions which represents while 55% of K99 awards were made to 226 lower-funded institutions (Fig. 5). R00 awards were distributed slightly more evenly with only 32% of R00 awards going to the top 25 highest-funded institutions and 68% going to lower funded institutions which represents 327 institutions. This means that the top 25 highest funded institutions had an average of 65.5 K99 awards and 38.2 R00 awards per institution. In contrast, the lower-funded institutions had an average of 8.4 K99 awards and 5.8 R00 awards per institution. While a greater percentage of both K99 and R00 awards are made at lower-funded public and private institutions, the top 25 highest funded institutions clearly receive and retain a far greater share of both K99 and R00 awardees.

Revised version: The majority of K99 awardees were at private institutions with 28% from the top 25 highest funded institutions and 35% from other private institutions. The remaining 37% of K99 awardees were from public institutions (Fig. 5A and Tables 3-4). We examined the composition of institutions represented by K99 awards (Fig. 5B). We found that 45% of K99 awards were made to the top 25 highest funded institutions which represents only 24 institutions while 55% of K99 awards were made to 226 lower-funded institutions (Fig. 5B). The percent of K99 awards made at each institution type has remained consistent on an annual basis (Fig. 5C). The majority of R00 awardees were at lower-funded institutions with 38% at public institutions 38% and 29% at private (Fig. 5D and Tables 3-4). Only 32% of R00 awards were made to the top 25 highest NIH-funded institutions. R00 awards were distributed slightly more evenly amongst institutions with 32% of R00 awards going to the top 25 highest-funded institutions and 68% going to lower funded institutions which represents 327 institutions (Fig. 5E). This means that the top 25 highest funded institutions had an average of 65.5 K99 awards and 38.2 R00 awards per institution. In contrast, the lower-funded institutions had an average of 8.4 K99 awards and 5.8 R00 awards per institution. The percent of R00 awards made at each institution type was also consistent on an annual basis (Fig. 5F). While a greater percentage of both K99 and R00 awards are made at lower-funded public and private institutions, the top 25 highest funded institutions clearly receive and retain a far greater share of both K99 and R00 awardees.

Earlier version: The percentage of K99 and R00 recipients by the institution type and the number of each institution type that has received K99 and R00 awards from 2006-2022.

Revised version: The percentage of recipients by the institution type for K99 (A) and R00 (B) and the number of different institutions within institution type that has received K99 (B) and R00 (E) awards from 2006-2022. The percentage of institution types receiving K99 (C) and R00 (F) awards by year from 2006-2022.

Being female, staying at the same K99 and R00 institution, and activating the R00 at a lower-funded institution are disadvantages to future funding success

We sought to determine which factors impacted whether K99/R00 awardees received major subsequent extramural funding in the form of an NIH R01, DP2, or R35 MIRA grant. We refer to these as “major awards” as they fund a similar total cost, timeframe, and hold similar weight for tenure and promotion considerations. Thus, we calculated a Cox proportional hazard model which includes whether the K99/R00 awardees stayed at the K99 institution for the R00 phase or left, were men or women, whether the K99 or R00 institutions were private or public, and whether the K99 and R00 institutions were one of the top 25 institutions for NIH funding in 2022 (Fig. 6). Individuals that stayed at the same institution for the K99 and R00 phase were at the greatest disadvantage for receiving future funding (HR: 0.7095, P<0.0001). The next strongest factor in receiving future R01 funding was whether the R00 institution was in the top 25 largest NIH funding portfolios (P<0.0001). The third factor was whether the candidate was a man or woman (P=0.001), consistent with previous findings that women are at a disadvantage in academic career placement.34 The NIH funding portfolio of the K99 institution and whether the K99 institution was private or public did not impact success in receiving future R01 funding.

Added in the revised version: The similarity in effect sizes suggests a cumulative advantage, or disadvantage, model. K99/R00 awardees that were men and moved to a new institution with large NIH funding portfolios for the R00 phase of the award have the shortest median time to receiving a subsequent major NIH award (median 4.6 years from the R00 budget start date to major award budget start date) as well as greater overall percent chance of ever receiving a major award (84.6%; Fig. 6). Each cumulative disadvantage increased the median time to receive a major NIH award by about 6 months and reduced the overall chance to receive subsequent major NIH award by 5-6%. The median time to receive subsequent major NIH awards for K99/R00 awardees that were women that stayed at an R00 institution with smaller NIH funding portfolios is 7.4 years and there is only a 60% chance to receive any major NIH award. This amounts to 2.8 years longer and 24.6% lower chance to receive any subsequent major funding compared with men who move to institutions with the largest NIH funding portfolios for the R00 phase of the award. This length of time is critical for tenure-track faculty who typically have between 6-7 years to earn tenure which often hinges on receipt of a major research grant for research-intensive faculty.

Cox proportional hazard model examining K99/R00 awardees success at receiving subsequent major extramural NIH awards by mobility, classified as women or men, and R00 institutional NIH funding level. Mobility is classified whether the K99 awardee moved to a new institution for the R00 award activation or stayed at the same institution. Classification of women or men was determined by name (see Methods for description and limitations). R00 institution funding level is classified as: 1 = 25 highest NIH funded institutions in 2022 and 2: all other institutions. A) Survival curve demonstrating all 8 possible classifications. B) Cox proportional hazard model hazard ratio, 95% confidence interval, and P-value. Survival plots showing the individual components of the full model including C) NIH funding + mobility, D) mobility + women/men, E) NIH funding + women/men, F) mobility, G) women/men, and H) NIH funding.

Removed in the revised version: The similarity in effect sizes suggests a cumulative advantage, or disadvantage, model. K99/R00 awardees that were men and moved to a new institution with large NIH funding portfolios for the R00 phase of the award have the shortest median time to receiving a subsequent major NIH award (median 4.6 years from the R00 budget start date to major award budget start date) as well as greater overall percent chance of ever receiving a major award (84.6%; Fig. 6). Each cumulative disadvantage increased the median time to receive a major NIH award by about 6 months and reduced the overall chance to receive subsequent major NIH award by 5-6%. The median time to receive subsequent major NIH awards for K99/R00 awardees that were women that stayed at an R00 institution with smaller NIH funding portfolios is 7.4 years and there is only a 60% chance to receive any major NIH award. This amounts to 2.8 years longer and 24.6% lower chance to receive any subsequent major funding compared with men who move to institutions with the largest NIH funding portfolios for the R00 phase of the award. This length of time is critical for tenure-track faculty who typically have between 6-7 years to earn tenure which often hinges on receipt of a major research grant for research-intensive faculty.

Faculty at HBCUs mostly have doctoral degrees from HBCUs and public and private institutions outside of the top 25 NIH funding category

It is clear that diversity in all metrics is beneficial for scientific progress. We sought to examine the rates of funding for those at historically under-funded institutions. We examined HBCUs in NIH Reporter, we found that there were 8 active R01 awards at Morehouse School of Medicine, 9 at Howard, 5 at Meharry Medical College, and 5 at Jackson State University in April 2023. Strikingly, of all 2703 K99 awardees that transitioned to R00 awards, no K99 awardees were from HBCUs, and only 2 K99 awardees activated R00 awards at HBCUs. Thus, there are not enough data available to determine how the classification of HBCU impacts future R01 funding success; however, Wapman and colleagues have made an extensive dataset of all United States faculty from 2010 to 2020 and where they received their doctoral degrees publicly available.15 Wapman et al., did not demonstrate doctoral degrees or faculty hiring practices at HBCUs. This prompted us to examine faculty flow at HBCUs. We examine where faculty with doctoral degrees from HBCUs ended up for their faculty career (Fig. 7A) and where faculty members at HBCUs received their doctoral degrees (Fig. 7B) specifically related to biological and biomedical faculty. We included only biology, biomedical, and health related fields (see Methods) which excluded fields such as computer science, business, language, etc., which resulted in a total of 65,120 faculty flows from doctoral degree programs.

A) Sankey diagram of faculty members that received doctoral degrees from HBCUs (on left) and the faculty institution they moved to (on right). B) Sankey diagram of faculty members at HBCUs (on right) and where they received their doctoral degrees (on left). Interactive versions of these figures can be found at: https://dantyrr.github.io/K99-R00-analysis/HBCU_deg.html for Figure 7A and https://dantyrr.github.io/K99-R00-analysis/HBCU_faculty.html for Figure 7B.

Out of the 65,120 total faculty flows, there are data on 296 faculty members with doctoral degrees from HBCUs with a known doctoral degree granting institution. There were 529 faculty at HBCUs with a known doctoral degree granting institution. Those with doctorates from primarily-underrepresented minority-serving institutions are less likely to be in academic positions.35 The HBCU granting the most doctoral degrees that went on to receive faculty appointments in the United States were from Howard University with 116 individuals. The majority of these are working as faculty members at Howard (56). The HBCU with the next most doctoral degrees in faculty positions is Florida A&M with 44 and the bulk of these (32) are employed at Florida A&M. Most of the institutions with faculty that have doctoral degrees from HBCUs are other HBCUs and public and private institutions. Notably there are very few HBCU doctoral degree holders working at public or private institutions within the top 25 NIH funding category (Fig. 7A). When examining faculty flows to HBCU’s it is apparent that the HBCU faculty ranks employ a large number of faculty with HBCU doctoral degrees (Fig. 7B). There are more faculty at Howard University with doctoral degrees from a range of different types of institutions which are primarily public and private and institutions outside of the top 25 NIH funding category; however, there are some (9 from University of Michigan, 1 from Vanderbilt, 4 from University of Wisconsin, and 3 from UNC). For the other HBCUs, the faculty ranks are dominated primarily by faculty with doctoral degrees from other HBCUs and public universities, primarily outside of the top 25 NIH funding category. These data highlight stark differences in hiring practices between HBCUs and institutions that hire significant numbers of R00 awardees. We quantified where faculty at a subset of institutions received their doctoral degrees with annotations for the top 15 institutions (Supplemental Fig. S1).

Self-hiring and regional hiring of faculty is ubiquitous

Some institutions tend to hire more faculty with doctoral degrees from the same institution (i.e., self-hires). A quarter of faculty at USC have USC doctoral degrees (Supplemental Fig. S1). In contrast, only ∼10-12% of faculty at Duke and Yale have doctoral degrees from Duke and Yale (Supplemental Fig. S1). In contrast, at Princeton, there are relatively few faculty with Princeton doctoral degrees. Certain institutions (Harvard, Columbia, U Penn, Stanford, Johns Hopkins, University of Michigan, UC Berkeley, Duke, and UNC) are nearly ubiquitous in the top 15 doctoral degrees for faculty at many institutions. There is also regional bias where institutions in the same region or city are more heavily represented. For example, doctoral degree holders from Tufts and Boston University are common for Harvard Faculty. A significant number of faculty at University of Michigan have doctoral degrees from Wayne State University, which is in Detroit, Michigan. The faculty at southeastern institutions have greater representation from Vanderbilt, Emory, Duke, University of North Carolina, and Washington University in St. Louis compared with the rest of the country. Similarly, the University of California system institutions have greater representations from other UC system affiliates including UC Berkeley, UCSD, UCSF, UCLA, USC, and Caltech (Supplemental Fig. S1).

Institution faculty hiring is consistent with institution type

We also examined the type of institutions that faculty received their doctoral degrees at (see Fig. 3). Massachusetts General Hospital, Yale, and Stanford are private institutions with large NIH funding portfolios along with public institutions like UCSF and U of Michigan. We compare these with historically Black colleges and universities and foreign institutions. The faculty at private institutions with large NIH funding portfolios tend to have degrees from other well-funded private institutions (Supplemental Fig. S2). The faculty at public institutions with large NIH funding portfolios (i.e., UCSF and University of Michigan) tend to have doctoral degrees from other public institutions with large NIH funding portfolios. Similarly, faculty at public institutions with smaller NIH funding portfolios have degrees from other public institutions with smaller NIH funding portfolios (Supplemental Fig. S2). The faculty at HBCU’s tend to have doctoral degrees from public institutions with smaller NIH funding portfolios as well as other HBCUs (Supplemental Fig. S2). Faculty members with doctoral degrees from HBCUs are rare and do not exceed more than 2% of the total faculty except at other HBCUs and Middle Tennessee State University which has 3.9%. In general, the practice of self-hiring contributes to institutions having faculty with the same doctoral degree-type as the institution they are employed at.

Discussion

It has previously been shown that a minority of institutions populate the faculty ranks of the bulk of other institutions along prestige hierarchies across fields and disciplines.7,15 These findings are in-line with the Matthew effect whereby success early on increases the chance of success in the future.36,37 Our analyses focus on the K99/R00 career development award. We demonstrate there are significant disadvantages associated with mobility, gender identity, type of R00 institution, and whether faculty candidates move from the K99 institution or stay at the same institution. In the context of the Matthew effect, there are significant advantages for receiving the K99 award at a highly funded institution or for being classified as a man or having mobility in the choice of an R00 institution. Scientists are often the most mobile at the earliest stages of their careers whether by choice or necessity.38 Academic mobility, especially in early stage investigators has been described as “coerced movement from job to job” and has also been termed “forced mobility”.39 International post-doctoral fellows may have a cultural reasons for moving to the United States based in perceived bias regarding the United States’ scientific training.40 Our analysis also demonstrates that K99 to R00 faculty flow generally follows similar hierarchies with some deviation for specific institutions. We find that a minority of private and public institutions with the largest NIH funding portfolios produce the majority of K99 awardees, and there are a handful of private and public institutions with the largest NIH funding portfolios that recruit the most R00 awardees. Many K99 awardees from institutions with the largest NIH funding portfolios move to public institutions and institutions with smaller NIH funding portfolios for the R00 phase of the award. These also happen to be less concentrated on the East and West coast. This supports the universal core-periphery structure of hiring that Wapman and colleagues conceptualized but suggests that prestige hierarchies in academic biomedical science align with NIH funding portfolio size.15 Despite this, there are also K99/R00 awardees that move up the prestige hierarchy to public and private institutions with the largest NIH funding portfolios and some individuals that stay within the same category for both K99 and R00 award phase. There are stark differences in future major NIH grant funding depending on several factors including NIH funding portfolio size of R00 institution, gender identity, and whether the awardee moved or stayed at the K99 institution.

Notably, whether the K99 institution was private or public or the size of the NIH funding portfolio did not impact future funding success except when considering whether the candidate eventually moved for the R00 phase of the award. One key component of activating the R00 award is to demonstrate independence from the candidate’s postdoctoral advisor because the candidate will need to establish independence from the postdoctoral advisor to be competitive to receive future NIH funding. Moving to a new institution clearly establishes independence; however, there are many pressures to stay at the K99 institution including familiarity with the institution, location, and family. K99/R00 awardees that stay at the same institution have an increase in the median time to future major NIH funding (either R01, DP2, or R35 MIRA) and have a reduced overall chance to receive any major NIH award. This may be due to reduced start-up packages for internal candidates, which could leave them less well-suited to perform the necessary research to be competitive for R01 funding. Candidates that remain at the same institution may also have additional duties in service work that drain their time. They may have more work on committees given the knowledge, experience, familiarity, and institutional history. Alternatively, this may be an internal bias of candidates within the institution or of grant reviewers who do not view the candidate as truly independent from their postdoctoral advisor. There are also cultural biases that largely reward white and Asian male academics by perceiving them to be more productive and de-valuing work done by non-white and non-Asian men.41,42 These data bring to light several factors that affect the scientific human experience. Moving institutions throughout one’s academic career can cause significant hardship on anyone. First generation scientists and those in systematically marginalized groups may lack support systems that those from non-marginalized groups have. Familial/friendships ties play important roles in the happiness, support, and productivity in the lives of scientists. Compounding the burden of first-generation scientists and systemically marginalized scientists with the burden of moving away from supports offered by family and friends may have a significant negative impact on those scientists. The second is that obtaining K99 funding and staying at the same institution comes with inherent biases within the system designed to support scientific talent. How these two factors are at odds with each other is of significant note to the wider scientific community. Ameliorating these disparities to increase equity should be of future focus.

Earlier version: Additional disadvantages to future funding success of K99/R00 awardees is whether the candidate is a man or a woman and whether the candidate is from an institution with a large NIH funding portfolio. Each of these factors can prolong the median time to future funding by about half a year and reduces the overall chance of securing future funding by ∼5.5%. These cumulative disadvantages may contribute to lower likelihood of receiving promotion and tenure and may contribute to the gender disparities in faculty ranks and in leadership positions. The United States has a more even gender balance for tenure-track positions than some other countries including Japan and South Korea43; however, it is apparent that women are at a disadvantage compared with men despite this. There may also be additional disadvantages for K99/R00 awardees who identify as transgender, gender non-binary, or non-conforming, or for individuals from systematically marginalized groups, or those with disabilities. There have been recommendations to promote a sustainable biomedical research enterprise.44 Future work is needed to examine how intersectional identities impact scientific success within the academic workforce.

Revised version: Additional disadvantages to future funding success of K99/R00 awardees is whether the candidate is a man or a woman and whether the candidate is from an institution with a large NIH funding portfolio. Each of these factors can prolong the median time to future funding by about half a year and reduces the overall chance of securing future funding by ∼5.5%. These cumulative disadvantages may contribute to lower likelihood of receiving promotion and tenure and may contribute to the gender disparities in faculty ranks and in leadership positions. The United States has a more even gender balance for tenure-track positions than some other countries including Japan and South Korea43; however, it is apparent that women are at a disadvantage compared with men despite this. There may also be additional disadvantages for K99/R00 awardees who identify as transgender, gender non-binary, or non-conforming, or for individuals from systematically marginalized groups, or those with disabilities. There have been recommendations to promote a sustainable biomedical research enterprise.44 Future work is needed to examine how intersectional inequalities impact scientific success within the academic workforce.

We can speculate why K99 awards most often go to individuals at institutions with the largest NIH funding portfolios; however, data on individual K99/R00 applications that were not funded are not publicly available. Our data show that 772 scientists with funded K99 awards did not get R00 awards and 302 of these were from 2019 or earlier. Those made from 2020 onward may still have the possibility to convert to the R00 phase, even after the 2-year K99 period with a 1-year no-cost extension. The factors that prevented the other 302 K99 awardees from 2019 and earlier unable to convert their K99-R00 grants is cause for concern within our greater academic community. Possible explanations include leaving the biomedical workforce, accepting tenure-track positions or other positions abroad, or by simply not successfully securing a tenable tenure-track offer. Of the funded K99/R00 awardees, we can make the assumption that candidates at institutions with large biomedical research portfolios and built-in support to write competitive grant applications that may be lacking at institutions with smaller research enterprises. Institutions with the largest NIH funding portfolios may facilitate greater access to making scientific discoveries. Movement of postdocs from these institutions to independent faculty positions at institutions with the smaller NIH funding portfolios, may increase the effectiveness of the faculty at those institutions. We must also consider that senior faculty themselves review and score K99/R00 award applications, and by doing so, determine which candidates ultimately receive these awards. A contrary perspective is that K99/R00 awardees at institutions perceived as more prestigious receive their awards in part because of the name and reputation of their institution. Thus, the interpretation of our data is that the system of reviewing and scoring grants has inherent bias towards less prestigious institutions which is currently being reviewed at the NIH. We can extend this speculation of bias to explain why men who receive R00-awards and move to institutions with the largest NIH funding portfolios are the most successful at securing subsequent major NIH awards in contrast to women, people who stay at the same institution for the R00 phase, and individuals at institutions with smaller NIH funding portfolios.

Of note, no K99 award has been made to a candidate at an HBCU, and only 2 awardees activated R00 awards at HBCUs. Most HBCUs specialize in undergraduate education; however, several perform biomedical research. Thus, we used a different dataset of faculty institutions and doctoral degree institutions to examine how faculty hiring practices differ between institution types, in particular, HBCUs. In agreement with others, we also found that self-hiring is common. 45 The self-hiring rate generally between 10-20% most institutions including HBCUs. One stark difference between HBCUs and all other US research institutions is that HBCUs appear to be the only institution type that hires faculty with doctoral degrees from other HBCUs. For example, of the 938 biomedical faculty members at Harvard, only 1 has a doctoral degree from an HBCU. Some institutions have lower rates of self-hiring. For example, Princeton faculty have degrees primarily from other Ivy league institutions but doctoral degrees from Princeton only represent the 14th most common for biomedical sciences. In agreement with the known and described systemic inequalities within the academic system we have highlighted in this article, self-hiring within HBCUs may occurs to enable survival of these institutions. Although many HBCUs focus primarily on undergraduate education, several have large biomedical research enterprises, and it is shocking that no K99 awardees have ever been from HBCUs. The K99 MOSAIC program, established in 2020 may increase equity; however, an effort to fund career development awards specifically at HBCUs may also increase equity in this area.

We do not know how many un-funded K99 applications may have originated from HBCUs. A freedom of information act request in 2018 by Dr. Pickett (request # 47950) for fiscal years 2007-2017 for K99/R00 awardees found that the successful award rate for white K99 applicants was 31.0% which was higher than both Asian (26.7%) and Black (16.2%) applicants. Total successful R00 transition rates were also higher among white K99 applicants (77%) compared to Asian (76.1%) and Black (60.0%) awardees.46 Although data on specific applicants is not available to determine trajectories from specific institution types, these data demonstrate that while improvements in equity to male and female applicants have improved; there is still a significant racial equity gap.

These analyses exclude those that received a K99 but were either unable to transition to an R00 award for various reasons including receiving a faculty position in the United States without needing the R00, receiving a faculty position outside of the United States, pursuing another line of work, or other unknown reasons. Furthermore, these analyses are limited to academic institutions in the United States and do not examine academic centers in other countries. When interpreting these data, it is important to consider that each data point corresponds to an individual with unique motivations for choice of postdoctoral research advisor, institution, and laboratory as well as autonomy in whether to pursue a faculty position and where that may be. Complex consideration of start-up package, institutional clout, and individual motivations factor into this decision. In addition, the choice to give K99/R00 awards and subsequent NIH awards lies mostly with grant reviewers who are generally mid- to senior-career faculty that may have different levels of conscious or unconscious bias for or against various factors that contribute to the overall score given. This dataset does not account for the motivation behind these complex decisions, and an understanding of these motivating factors would be enlightening.

It is important to note the limitations of this study. The present study is focuses on K99/R00 awardees and is limited to faculty flows within the United States. While similar practices may exist in other countries, there are likely significant differences also. Another limitation due to our focus on K99/R00 awardees, is that scientists without K99/R00 awards are excluded from our analyses. Further, the public data on K99/R00 awardees does not provide demographic details including self-identified gender, race, ethnicity, age, and degree type. The gender identity used here was assigned as a binary man or woman gender based on the first name of the candidate. If these demographic data were available, a critical analysis of bias and how these impact outcomes could be made. For our analysis of K99 awardees, we do not have information on the doctoral degree granting institution, and for our exploration of where faculty members at HBCUs received their doctoral degrees, we do not have information on their postdoctoral research experience. If we had these details, we could make a more complete examination of flow from doctoral degree institution to postdoctoral institution to faculty institution and then to determine outcomes based on each of these. Future research should aim to complete a more comprehensive assessment of faculty flows and outcomes.

Here we demonstrate the complex nature of faculty hiring within the NIH grant system that exists today in the biomedical sciences and related fields. The K99/R00 award mechanism undoubtedly increases awardees’ chance of securing a faculty position. As one of the most coveted NIH grants for postdocs, we have examined the flow of K99 awardees to R00 institutions and how these flows impact future NIH award funding. Many factors besides those considered and quantified here contribute to where K99 awardees choose to begin their faculty careers, and the K99 award itself does not guarantee securing a faculty position. Despite this, we have identified factors that pose significant disadvantages to future funding success for K99/R00 awardees which likely influence funding success more broadly. Future work must examine the role of ethnic and racial bias in these domains. As the K99 MOSAIC program becomes more established, a comparison of this program and the K99 grant mechanism explored here may reveal whether the MOSAIC strategy is effective at promoting equity for under-represented minorities in the biomedical faculty ranks. By quantifying and understanding these factors, grant reviewers, faculty hiring committees, department chairs, and funding bodies may be able to more equitably award and administer grants and evaluate faculty candidates. Improvements to the faculty hiring process can be achieved centrally through the NIH creating a more equitable grant funding system or within the many diverse faculty searches.47,48 There is only one NIH and many searches, so systemic change will likely lead to more impactful change; however, both options could be implemented simultaneously for the optimal benefit of the scientific community.

Supporting information